<a href="https://colab.research.google.com/github/petre001/Demo/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jeffrey Petrella"
COLLABORATORS = "None"
```

---

# Python for Data Science Bootcamp
## Assignment 4

NOTE: This assignment has 4 problems.  Please write your code in the designated portion of the autograded answer cell.

For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](https://github.com/Python-for-Data-Science-Bootcamp/a4-petre001/blob/master/jupyter.png?raw=1)

©2020 Luis Morales. All rights reserved. 

## Problem 1 (1 point)

Create a function that loads the file (WorldCups.csv), and then return the following DataFrame sorted by the numbers in the 'Score' column. For each world cup, the winner gets 4 points, the runners-up gets 2 points, the third place gets 1 point. 
![image.png](attachment:image.png)

<!-- grade_id: myfunc, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    worldcup = pd.read_csv('/Python-for-Data-Science-Bootcamp/a4-petre001/WorldCups.csv')
    print(worldcup.head())
    raise NotImplementedError()
```

```python
pwd
```

Output:
```
'/content'
```

<!-- grade_id: myfunctest, points: 1, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert(res.shape==(19, 4))
assert(res['Score'].values==np.array([28, 26, 21, 14,  8,  8,  7,  4,  4,  4,  4,  4,  2,  1,  1,  1,  1,
        1,  1])).all()
```

## Problem 2 (1 point)

Write a function that loads the WorldCupMatches.csv file, calculates and returns a DataFrame representing the top 5 teams descendingly ordered by the number of goals in a single world cup".
![image.png](attachment:image.png)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: myfunctiontest, points: 1, grade: true, solution: false -->
```python
res=myfunction()
assert res.shape==(5,1)
assert (res.values==np.array([[27.],[25.],[23.],[22.],[19.]])).all()
assert int(res.loc[1954,'Hungary'])==27
assert int(res.loc[1954,'Germany FR'])==25
assert int(res.loc[1958,'France'])==23
assert int(res.loc[1950,'Brazil'])==22
assert int(res.loc[1970,'Brazil'])==19
```

## Problem 3 (1 point)

Create a function that loads the files WorldCupMatches.csv and WorldCupPlayers.csv to calculate all the matches that have goal in the first minute, and output the following dataframe.
![image.png](attachment:image.png)

<!-- grade_id: myfunction2, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: myfunction2test, points: 1, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert res.shape==(12,7)
assert res['Home Team Goals'].sum()==37
assert res['Away Team Goals'].sum()==20
```

## Problem 4 (1 point)

Given a dataframe as follows:

full_monte = pd.DataFrame(data=['Graham Chapman', 'John Cleese', 'Terry Gilliam'],columns=['Name'])
![image.png](attachment:image.png)
return a new dataframe that has an additional column of last name, calculated from the Name column as follows:

<!-- grade_id: prob4, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    full_monte = pd.DataFrame(data=['Graham Chapman', 'John Cleese', 'Terry Gilliam'],columns=['Name'])
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: prob4test, points: 1, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
full_monte2=myfunction()
assert full_monte2['Last Name'].iloc[0]=='Chapman'
assert full_monte2['Last Name'].iloc[1]=='Cleese'
assert full_monte2['Last Name'].iloc[2]=='Gilliam'
```